---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: d98503b49c
xp: 100
```

print() is a output statement in python. To display message you can use print("message") 

`@instructions`
display welcome message by using print()

`@hint`
hello have typed welcome in " "

`@pre_exercise_code`
```{python}
print("--")
```

`@sample_code`
```{python}

```

`@solution`
```{python}
print("welcome")
```

`@sct`
```{python}

```
